<a href="https://colab.research.google.com/github/OmamaRawahi/O-M/blob/main/Copy_of_2_K_Anonymity.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# CS211: Data Privacy
## In-Class Exercises, week of 9/5/2022

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

```python
adult

```

Output:
```
Name         DOB          SSN    Zip  Age  \
0       Karrie Trusslove    9/7/1967  732-14-6110  64152   39   
1       Brandise Tripony    6/7/1988  150-19-2766  61523   50   
2          Brenn McNeely    8/6/1991  725-59-9860  95668   38   
3            Dorry Poter    4/6/2009  659-57-4974  25503   53   
4            Dick Honnan   9/16/1951  220-93-3811  75387   28   
...                  ...         ...          ...    ...  ...   
32556       Ardyce Golby  10/29/1961  212-61-8338  41328   27   
32557      Jean O'Connor   6/28/1952  737-32-2919  94735   40   
32558   Reuben Skrzynski    8/9/1966  314-48-0219  49628   58   
32559        Caye Biddle   5/19/1978  647-75-3550   8213   22   
32560  Hortense Hardesty   5/11/2012  690-42-5638  86153   52   

              Workclass  fnlwgt   Education  Education-Num  \
0             State-gov   77516   Bachelors             13   
1      Self-emp-not-inc   83311   Bachelors             13   
2               Private  215646     HS-grad          
```

## Question 1

Write code to group the `adult` dataset by a single column and count the number of members in each group.

<!-- grade_id: cell-055399c92e676783, points: 0, grade: false, solution: true -->
```python
def group_one_count(col):
    return adult.groupby(col).size()

```

```python
group_one_count('Education')
```

Output:
```
Education
10th              933
11th             1175
12th              433
1st-4th           168
5th-6th           333
7th-8th           646
9th               514
Assoc-acdm       1067
Assoc-voc        1382
Bachelors        5355
Doctorate         413
HS-grad         10501
Masters          1723
Preschool          51
Prof-school       576
Some-college     7291
dtype: int64
```

<!-- grade_id: cell-befc11c8be0bf91b, points: 5, grade: true, solution: false -->
```python
s = group_one_count('Education')
assert s['10th'] == 933
assert s['9th'] == 514
assert s['Some-college'] == 7291
```

## Question 2

Write code to group the `adult` dataset by two columns and count the number of members in each group.

<!-- grade_id: cell-f6a37efe405c4a98, points: 0, grade: false, solution: true -->
```python
def group_two_count(col1, col2):
     return adult.groupby([col1,col2]).size()
```

```python
group_two_count('Occupation', 'Education')
```

Output:
```
Occupation        Education   
Adm-clerical      10th             38
                  11th             67
                  12th             38
                  5th-6th           6
                  7th-8th          11
                                 ... 
Transport-moving  Doctorate         1
                  HS-grad         825
                  Masters          10
                  Prof-school       3
                  Some-college    283
Length: 201, dtype: int64
```

<!-- grade_id: cell-88caedc5d6e44568, points: 5, grade: true, solution: false -->
```python
s = group_two_count('Occupation', 'Education')
assert s['Transport-moving']['Doctorate'] == 1
assert s['Adm-clerical']['10th'] == 38
```

## Question 3

Write code to perform a differencing attack to determine Brenn McNeely's age using only the `mean` aggregation function over large groups.

<!-- grade_id: cell-6fb63040052b0857, points: 0, grade: false, solution: true -->
```python
def get_brenns_age():
    a1=(adult[adult['Name']!='Brenn McNeely']['Age'].mean())*len(adult[adult['Name']!='Brenn McNeely']['Age'])
    a2=(adult['Age'].mean())*len(adult['Age'])
    return a2-a1
brenns_age = get_brenns_age()
```

<!-- grade_id: cell-d66b7d6c32d256c6, points: 5, grade: true, solution: false -->
```python
assert brenns_age == 38.0
```

```python
df = adult[['Education', 'Marital Status', 'Target']]
df.head(10)
```

Output:
```
Education         Marital Status Target
0  Bachelors          Never-married  <=50K
1  Bachelors     Married-civ-spouse  <=50K
2    HS-grad               Divorced  <=50K
3       11th     Married-civ-spouse  <=50K
4  Bachelors     Married-civ-spouse  <=50K
5    Masters     Married-civ-spouse  <=50K
6        9th  Married-spouse-absent  <=50K
7    HS-grad     Married-civ-spouse   >50K
8    Masters          Never-married   >50K
9  Bachelors     Married-civ-spouse   >50K
```

## Question 4

Does the dataset `df` satisfy $k$-Anonymity for $k=3$? Why not?

<!-- grade_id: cell-a9793d100ba20894, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 5

What columns should we designate as quasi-identifiers in the dataset `df`, and why?

<!-- grade_id: cell-5cfd666564429892, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 6

Imagine the column `Target` is *not* a quasi-identifier, and we generalize the dataset to achieve $k$-Anonymity for $k=2$ as follows:

- Replace each education level below "HS Grad" with `< HS` and others with `>= HS`
- Replace marital status with `Married` and `Not Married`
- Delete rows as required to achieve $k$-anonymity for $k=2$

For which rows is a homogeneity attack possible, and why?

<!-- grade_id: cell-5932f02f09b5cacb, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 7

Write code to generalize the `Zip` column in the adult dataset.

<!-- grade_id: cell-be9a4264a0e58938, points: 0, grade: false, solution: true -->
```python


def generalize_zip(zip, l):
  # Generate a string of zeros of length l
  zeros = '0' * l
  zip_str = str(zip)
  # Concatenate
  coc = zip_str[:len(zip_str)-l] + zeros
  return int(coc)


```

<!-- grade_id: cell-256688fe78b8f7bf, points: 5, grade: true, solution: false -->
```python
# Test cases
assert generalize_zip(47401, 0) == 47401
assert generalize_zip(47401, 1) == 47400
assert generalize_zip(47401, 2) == 47400
assert generalize_zip(47401, 3) == 47000
assert generalize_zip(47401, 4) == 40000
assert generalize_zip(47401, 5) == 0
```