<a href="https://colab.research.google.com/github/OmamaRawahi/O-M/blob/main/Another_copy_of_Linking_Def_Attacks.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Data Privacy

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
adult_pii = adult[['Name', 'DOB', 'SSN', 'Zip', 'Age']]
adult_deid = adult.drop(columns=['Name', 'SSN'])
```

```python
adult
```

## Question 1 (20 points)

Using the dataframes `adult_pii` and `adult_deid`, write code to conduct a linking attack to recover the names of as many individuals in `adult_deid` as possible. Your solution should be parameterized by the set of columns to use in the attack.

<!-- grade_id: cell-e6389ceace711c32, points: 0, grade: false, solution: true -->
```python
def linking_attack(cols):
    return pd.merge(adult_pii[cols],adult_deid[cols])
```

<!-- grade_id: cell-251b9f1c6f7b4630, points: 20, grade: true, solution: false -->
```python
# TEST CASES for Question 1

assert len(linking_attack(['Zip'])) == 43191
assert len(linking_attack(['Zip', 'DOB'])) == 32563
assert len(linking_attack(['Zip', 'Age'])) == 32755
```

## Question 2 (5 points)

How many individuals in this dataset are uniquely identified by their Zip code? How many are uniquely identified by their age?

Hint: note that the number of *unique ZIP codes* is **different** from the number of *individuals uniquely identified by ZIP code*.

Hint: you can use the `value_counts` method to count the number of occurences of each value in a series.

<!-- grade_id: cell-0e380ac68dd0c77e, points: 0, grade: false, solution: true -->
```python
def unique_zipcode():
    a= adult['Zip'].value_counts()
    num= len(a[a == 1])
    return num

def unique_dob():
   a=adult['DOB'].value_counts()
   num= len(a[a == 1])
   return num

```

<!-- grade_id: cell-9b009427c61dbb2b, points: 5, grade: true, solution: false -->
```python
# TEST CASES (INCOMPLETE) for Question 2

assert unique_zipcode() > 0
assert unique_dob() > 0

```

## Question 3 (10 points)

Write code to determine the `Education-Num` of the individual named Ardyce Golby by performing a differencing attack. Your code should *only* use aggregate data to find Ardyce's education number.

<!-- grade_id: cell-ef28e136429ae8c2, points: 0, grade: false, solution: true -->
```python
def ardyce_education():
     a=adult[adult['Name'] != 'Ardyce Golby']['Education-Num'].sum()
     #we want the eduction-num
     b=adult['Education-Num'].sum()
     return b-a
```

<!-- grade_id: cell-d06fb722893a39a5, points: 10, grade: true, solution: false -->
```python
# TEST CASE for Question 3
assert ardyce_education() == 12
```